
# Assignment No 2
----
## Goal

Your goal in this assigment is to implement a Decision Tree Classifier.

**Note** Please note that you are allowed to use only those libraries which we have discussed in the class, i.e. numpy, scipy, pandas. You can use scikit-learn only for data splitting.

## Submission Instructions
You are required to submit the original notebook file on the Classsroom (with .ipynb extension), with complete set of outputs. Students failing to do so will get zero marks. 

*Please read each step carefully and understand it fully before proceeding with code writing*

## Plagiarism
Any form of plagiarism will not be tolerated and result in 0 marks.



### Decision Tree Classifier

Now in this assignment we will be implementing the Decision Classifier for Continuous attributes.

You will be testing your implementations with only one split criteria, namely:
 - Information Gain


```python
%pylab inline
import scipy.stats
import pandas as pd
from sklearn import datasets
from collections import defaultdict  # default dictionary 
```

Output:
```
Populating the interactive namespace from numpy and matplotlib
```

<!-- grade_id: node, points: 0, grade: false, solution: true -->
```python
class Node:
    def __init__(self, purity, klasslabel='', score=0, split=[], fidx=-1):
        self.lchild=None       
        self.rchild=None
        self.klasslabel=klasslabel        
        self.split=split
        self.score=score
        self.fidx=fidx
        self.purity=purity
        
    def set_childs(self,lchild,rchild):
        # YOUR CODE HERE
        self.lchild=lchild
        self.rchild=rchild
        
    def isleaf(self):
        # YOUR CODE HERE
        if self.lchild==None and self.rchild==None:
            return True
        else:
            return False
        
        raise NotImplementedError()
    def isless_than_eq(self, X):
        # YOUR CODE HERE
        raise NotImplementedError()
    def get_str(self):        
        if self.isleaf():
            return 'C(class={},Purity={})'.format(self.klasslabel,self.purity)
        else:
            return 'I(Fidx={},Score={},Split={})'.format(self.fidx,self.score,self.split)
    
```

```python
# A placeholder class 
# TODO: You have to implement the following class, remember from the lectures that you will 
# need to build a model for each different class you are trying to identify...
```

<!-- grade_id: tree, points: 0, grade: false, solution: true -->
```python
## Your code goes here...
# You might need to define auxliary classes for composition.. ?
class DecisionTree:
    ''' Implements the Decision Tree For Classification... '''
    def __init__(self, purityp, exthreshold, maxdepth=10):        
        self.purity=purityp
        self.exthreshold=exthreshold
        self.maxdepth=maxdepth
        pass
    
    def train(self, X, Y):
        ''' Train Decision Tree using the given 
            X [m x d] data matrix and Y labels matrix
            
            Input:
            ------
            X: [m x d] a data matrix of m d-dimensional examples.
            Y: [m x 1] a label vector.
            
            Returns:
            -----------
            Nothing
            '''
        nexamples,nfeatures=X.shape
        ## now go and train a model for each class...
        # YOUR CODE HERE
        depth=0   
        self.tree = self.build_tree(X,Y,0)
        
        
    def build_tree(self, X, Y, depth):
        # YOUR CODE HERE
        pass
        
    def test(self, X):
        
        ''' Test the trained classifiers on the given set of examples 
        
                   
            Input:
            ------
            X: [m x d] a data matrix of m d-dimensional test examples.
           
            Returns:
            -----------
                pclass: the predicted class for each example, i.e. to which it belongs
        '''
        pclasses=[]
        nexamples, nfeatures=X.shape
        pass
    
    def evaluate_numerical_attribute(self,feat, Y):
        return split_point,Max_Gain,left_C,right_C
           
    def predict(self, X):
        pass
    
    def _predict(self,node, X):
   
        raise NotImplementedError()

    def __str__(self):
        
        return self.__print(self.tree)        
        
     
    def find_depth(self):
        return self._find_depth(self.tree)
    
    
    def _find_depth(self,node):
        if not node:
            return
        if node.isleaf():
            return 1
        else:
            return max(self._find_depth(node.lchild),self._find_depth(node.rchild))+1
    
    def __print(self,node,depth=0):
        
        ret = ""

        # Print right branch
        if node.rchild:
            ret += self.__print(node.rchild,depth+1)

        # Print own value
        ret += "\n" + ("    "*depth) + node.get_str()

        # Print left branch
        if node.lchild:
            ret += self.__print(node.lchild,depth+1)
        
        return ret
```

# In the first case here, you are asked to train on all the data using only two features and classes. 

```python
#load the data set
data = pd.read_csv('./iris.data')
data.columns=['SepalLength','SepalWidth','PetalLength','PetalWidth','Class']
print (data)
```

Output:
```
SepalLength  SepalWidth  PetalLength  PetalWidth           Class
0            4.9         3.0          1.4         0.2     Iris-setosa
1            4.7         3.2          1.3         0.2     Iris-setosa
2            4.6         3.1          1.5         0.2     Iris-setosa
3            5.0         3.6          1.4         0.2     Iris-setosa
4            5.4         3.9          1.7         0.4     Iris-setosa
5            4.6         3.4          1.4         0.3     Iris-setosa
6            5.0         3.4          1.5         0.2     Iris-setosa
7            4.4         2.9          1.4         0.2     Iris-setosa
8            4.9         3.1          1.5         0.1     Iris-setosa
9            5.4         3.7          1.5         0.2     Iris-setosa
10           4.8         3.4          1.6         0.2     Iris-setosa
11           4.8         3.0          1.4         0.1     Iris-setosa
12           4.3         3.0          1.1         0.1     Iris-setosa
13           5.8         
```

```python
# Get your data in matrix
X = np.asarray(data[['SepalLength','SepalWidth','PetalLength','PetalWidth']].dropna())
Y=np.asarray(data['Class'].dropna())
print " Data Set Dimensions=", X.shape, " True Class labels dimensions", Y.shape   

```

Output:
```
Data Set Dimensions= (149, 4)  True Class labels dimensions (149,)
```

```python

```

```python
# here data is being set so that there are only two classes
Y[ Y == 'Iris-virginica'] = 'Iris-versicolor'
```

```python

feat=[0,1] # because we are only using two features in the beginning
dt=DecisionTree(0.95,5,2)
feat=[0,1]
dt.classes=np.unique(Y)
dt.nclasses=len(np.unique(Y))
split,mingain,Xlidx,Xridx=dt.evaluate_numerical_attribute(X[:,0],Y)
```

```python
dt.classes=np.unique(Y)
dt.nclasses=len(np.unique(Y))
split,mingain,Xlidx,Xridx=dt.evaluate_numerical_attribute(X[:,0],Y)
```

<!-- grade_id: split_gain, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_almost_equal, assert_equal

assert_equal(split, 5.45)
assert_almost_equal(mingain, 0.38, places=1)
```

Output:
```
/home/shaoor/anaconda2/lib/python2.7/unittest/case.pyc in assertEqual(self, first, second, msg)
    511         """
    512         assertion_func = self._getAssertEqualityFunc(first, second)
--> 513         assertion_func(first, second, msg=msg)
    514 
    515     def assertNotEqual(self, first, second, msg=None):

/home/shaoor/anaconda2/lib/python2.7/unittest/case.pyc in _baseAssertEqual(self, first, second, msg)
    504             standardMsg = '%s != %s' % (safe_repr(first), safe_repr(second))
    505             msg = self._formatMessage(msg, standardMsg)
--> 506             raise self.failureException(msg)
    507 
    508     def assertEqual(self, first, second, msg=None):

AssertionError: 5.5 != 5.45
```

```python
print dt
```

Output:
```
C(class=Iris-versicolor,Purity=1.0)
I(Fidx=1,Score=-1.44269516019e-10,Split=0.6)
    C(class=Iris-setosa,Purity=1.0)
```

```python
print dt.find_depth()
```

Output:
```
2
```

```python
# Split your data into training and test-set...
Xtrain=[],Ytrain=[]
Xtest=[],Ytest=[]
print " Training Data Set Dimensions=", Xtrain.shape, "Training True Class labels dimensions", Ytrain.shape   
print " Test Data Set Dimensions=", Xtest.shape, "Test True Class labels dimensions", Ytest.shape   

```

Output:
```
Training Data Set Dimensions= (104, 4) Training True Class labels dimensions (104,)
 Test Data Set Dimensions= (45, 4) Test True Class labels dimensions (104,)
```

```python
# Lets train a Decision Tree Classifier on Petal Length and Width
feat=[0,1]
dt=DecisionTree(0.95,5)
dt.train(Xtrain[:,feat],Ytrain)
print dt
```

Output:
```
C(class=Iris-versicolor,Purity=1.0)
I(Fidx=0,Score=0.351952789971,Split=5.5)
        C(class=Iris-versicolor,Purity=1.0)
    I(Fidx=0,Score=0.271794230203,Split=5.7)
            C(class=Iris-setosa,Purity=0.961538461538)
        I(Fidx=1,Score=0.153049305507,Split=3.1)
            C(class=Iris-setosa,Purity=1.0)
```

```python
#Lets test it on the set of unseen examples...
feat=[0,1]
pclasses=dt.test(Xtest[:,feat])

```

```python
print pclasses
```

Output:
```
['Iris-setosa', 'Iris-setosa', 'Iris-versicolor', 'Iris-setosa', 'Iris-versicolor', 'Iris-setosa', 'Iris-setosa', 'Iris-setosa', 'Iris-setosa', 'Iris-setosa', 'Iris-setosa', 'Iris-setosa', 'Iris-setosa', 'Iris-setosa', 'Iris-setosa', 'Iris-setosa', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor', 'Iris-versicolor']
```

```python
Lets see how good we are doing...
==================
```

```python
#Lets see how good we are doing, by finding the accuracy on the test set..
print np.sum(pclasses==Ytest)
print "Accuracy = ", np.sum(pclasses==Ytest)/float(Ytest.shape[0])
```

Output:
```
42
Accuracy =  0.933333333333
```

<!-- grade_id: two_features, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_greater_equal
feat=[0,1]
dt=DecisionTree(0.95,5)
dt.train(Xtrain[:,feat],Ytrain)
pclasses=dt.test(Xtest[:,feat])
acc = np.sum(pclasses==Ytest)/float(Ytest.shape[0])
assert_greater_equal(acc, 0.90)
```

# Lets Train on all four features....

```python
#Lets Train on all four features....

# Lets train a Decision Tree Classifier on Petal Length and Width
feat=[0, 1, 2, 3]
dt=DecisionTree(0.95,5)
dt.train(Xtrain[:,feat],Ytrain)
pclasses=dt.test(Xtest[:,feat])
#Lets see how good we are doing, by finding the accuracy on the test set..
print np.sum(pclasses==Ytest)
print "Accuracy = ", np.sum(pclasses==Ytest)/float(Ytest.shape[0])
```

Output:
```
45
Accuracy =  1.0
```

# Lets Train on all four features and for all three classes....

```python
X=np.asarray(data[['SepalLength','SepalWidth','PetalLength','PetalWidth']].dropna())
Y=np.asarray(data['Class'].dropna())
print " Data Set Dimensions=", X.shape, " True Class labels dimensions", Y.shape
```

Output:
```
Data Set Dimensions= (149, 4)  True Class labels dimensions (149,)
```

```python
# Split your data into training and test-set... 
# see the documentation of split_data in tools for further information...
Xtrain=[],Ytrain=[]
Xtest=[],Ytest=[]

print " Training Data Set Dimensions=", Xtrain.shape, "Training True Class labels dimensions", Ytrain.shape   
print " Test Data Set Dimensions=", Xtest.shape, "Test True Class labels dimensions", Ytest.shape   

```

Output:
```
Training Data Set Dimensions= (104, 4) Training True Class labels dimensions (104,)
 Test Data Set Dimensions= (45, 4) Test True Class labels dimensions (104,)
```

```python
feat=[0, 1,2,3]
dt=DecisionTree(0.95,5)
good=[0,1]
dt.train(Xtrain[:,good],Ytrain)
pclasses=dt.test(Xtest[:,good])
#Lets see how good we are doing, by finding the accuracy on the test set..
print np.sum(pclasses==Ytest)
print "Accuracy = ", np.sum(pclasses==Ytest)/float(Ytest.shape[0])
```

Output:
```
42
Accuracy =  0.933333333333
```

<!-- grade_id: all_features, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_greater_equal

feat=[0, 1, 2, 3]
dt=DecisionTree(0.95,5)
good=[0,1]
dt.train(Xtrain[:,good],Ytrain)
pclasses=dt.te(Xtest[:,feat])
#Lets see how good we are doing, by finding the accuracy on the test set..
acc = np.sum(pclasses==Ytest)/float(Ytest.shape[0])
assert_greater_equal(acc, 0.90)

```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-14-c424f873e222> in <module>
----> 1 from nose.tools import assert_greater_equal
      2 
      3 feat=[0, 1, 2, 3]
      4 dt=DecisionTree(0.95,5)
      5 good=[0,1]

ModuleNotFoundError: No module named 'nose'
```

What can you conclude ?
====================
Please write your observation....



print("hello")